---
id: bd7158d8c442edefaeb5bd18
title: Build a Pack of Sand
challengeType: 21
dashedName: build-a-pack-of-sand
---

# --description--

Sandpack Project instructions

# --hints--

Test 1

```js
assert.fail();
```

# --seed--

## --seed-contents--

```py
const a = 1;
```

# --solutions--

```py
const b = 2;
```
